---
title: 'Chapter Title Here'
description: 'Chapter description goes here.'
---

## Example coding exercise

```yaml
type: NormalExercise
key: e8c1edbe67
lang: python
xp: 100
skills: 2
```

This is an example exercise. kh kjhgjkhg kjhgkjhg kjhg kj gkj

`@instructions`
lkj lkjm:lk jlkj lkh kljh lkjh kljhlkjlj klj hkljh klj hkl lkj hklj

`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}
print("awesome")
```

`@sct`
```{python}
Ex().has_output("awesome", pattern = False)
success_msg('Congratulations! Your path to awesomeness has began!')
```
